# CS414 - Numpy Theory
Instructions: 
- Start the kernel: In the menu bar, select Kernel -> Restart kernel.
- Complete all sections with the comment `YOUR CODE HERE`.
- Run all code blocks to check the implementation: In the menu bar, select Cell -> Run All.

```python
import numpy as np
```

## Problem 1: Vector Initialization
Initialize a vector of `n` elements, all of which are `0`.

Example:

- Input:
```
n = 5
```

- Output:
```
[0., 0., 0., 0., 0.]
```

<!-- grade_id: cell-02aa6d96c8c51dc9, points: 0, grade: false, solution: true -->
```python
def init_zeros(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_zeros(5)
```

<!-- grade_id: cell-ebe669293315df42, points: 5, grade: true, solution: false -->
```python
# grading cell 1
assert init_zeros(5).shape == (5, )
assert np.sum(abs(init_zeros(5) - np.array([0, 0, 0, 0, 0]))) == 0
```

<!-- grade_id: cell-0ef3e5e7d6fdb3b8, points: 5, grade: true, solution: false -->
```python
# grading cell 2

```

## Problem 2: Vector Initialization
Initialize a vector of `n` elements, all of which are `1`.

Example:

- Input:
```
n = 5
```

- Output:
```
[1, 1, 1, 1, 1]
```

<!-- grade_id: cell-a3d940e1087fc999, points: 0, grade: false, solution: true -->
```python
def init_ones(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_ones(5)
```

<!-- grade_id: cell-962507eb0e79c742, points: 5, grade: true, solution: false -->
```python
# grading cell 3
assert init_ones(5).shape == (5, )
assert np.sum(abs(init_ones(5) - np.array([1, 1, 1, 1, 1]))) == 0
```

<!-- grade_id: cell-e65c048faccd7122, points: 5, grade: true, solution: false -->
```python
# grading cell 4

```

## Problem 3: Vector Initialization
Initialize a vector of `n` elements, all of which are `x`.

Example:

- Input:
```
n = 5
x = 3.5
```

- Output:
```
[3.5, 3.5, 3.5, 3.5, 3.5]
```

<!-- grade_id: cell-0c08e0ad3fa2f88b, points: 0, grade: false, solution: true -->
```python
def init_all_x(n, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_all_x(5, 3.5)
```

<!-- grade_id: cell-d011edd31917dd88, points: 5, grade: true, solution: false -->
```python
# grading cell 5
assert init_all_x(5, 3.5).shape == (5, )
assert np.sum(abs(init_all_x(5, 3.5) - np.array([3.5, 3.5, 3.5, 3.5, 3.5]))) == 0
```

<!-- grade_id: cell-82e149db874d1e6b, points: 5, grade: true, solution: false -->
```python
# grading cell 6

```

## Problem 4: Vector Initialization
Initialize a vector of size `n` consisting of integers increasing from `0` to `n-1`.

Example:

- Input:
```
n = 5
```

- Output:
```
[0, 1, 2, 3, 4]
```

<!-- grade_id: cell-f57ae1d754273897, points: 0, grade: false, solution: true -->
```python
def init_range(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_range(5)
```

<!-- grade_id: cell-f603c404941edf2b, points: 5, grade: true, solution: false -->
```python
# grading cell 7
assert init_range(5).shape == (5, )
assert np.sum(abs(init_range(5) - np.array([0,1,2,3,4]))) == 0
```

<!-- grade_id: cell-9095933a8175e6df, points: 5, grade: true, solution: false -->
```python
# grading cell 8

```

## Problem 5: Vector Initialization
Initialize a vector of size `n` consisting of integers increasing from `x, x+1, ..., x+n-1` (a sequence starting from `x`).

Example:

- Input:
```
n = 5
x = 3.5
```

- Output:
```
[3.5, 4.5, 5.5, 6.5, 7.5]
```

<!-- grade_id: cell-c48f88fadbb49ac3, points: 0, grade: false, solution: true -->
```python
def init_range(n, x):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_range(5, 3.5)
```

<!-- grade_id: cell-3cfbf5f25a97488b, points: 5, grade: true, solution: false -->
```python
# grading cell 9
assert init_range(5, 3).shape == (5, )
assert np.sum(abs(init_range(5, 3) - np.array([3, 4, 5, 6, 7]))) == 0
```

<!-- grade_id: cell-6d2494de753e3271, points: 5, grade: true, solution: false -->
```python
# grading cell 10

```

## Problem 6: Matrix Initialization
Initialize a matrix of size `m x n` consisting of all elements as `K`.

Example:

- Input:
```
m = 3
n = 4
K = 5
```

- Output:
```
[[5., 5., 5., 5.],
 [5., 5., 5., 5.],
 [5., 5., 5., 5.]]
```

<!-- grade_id: cell-647a473a18b7e01f, points: 0, grade: false, solution: true -->
```python
def init_mat_K(m, n, K):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_mat_K(3, 4, 5)
```

Output:
```
<ipython-input-8-6ea02f736741> in <cell line: 2>()
      1 # testing with an example
----> 2 init_mat_K(3, 4, 5)

<ipython-input-7-cc49970cbb53> in init_mat_K(m, n, K)
      1 def init_mat_K(m, n, K):
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-fed0cc85f74a9f21, points: 5, grade: true, solution: false -->
```python
# grading cell 11
assert init_mat_K(2,3,4).shape == (2, 3)
assert np.sum(abs(init_mat_K(2,3,4) - np.array([[4, 4, 4], [4, 4, 4]]))) == 0
```

<!-- grade_id: cell-3e3a1d7e81c73d21, points: 5, grade: true, solution: false -->
```python
# grading cell 12

```

## Problem 7: Matrix Initialization
Initialize a radom matrix of size `m x n` consisting integers within a given range.

Example:

- Input:
```
m = 3
n = 4
min_val = 1
max_val = 10
```

- Output:
```
[[5, 9, 7, 2]
 [4, 3, 8, 6]
 [1, 7, 2, 9]]
```

```python
def init_mat_randints(m, n, min_val, max_val):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_mat_randints(3, 4, 1, 10)
```

```python
# grading cell 13
assert init_mat_randints(2, 3, 1, 10).shape == (2, 3)
assert np.all((init_mat_randints(2, 3, 1, 10) >= 1) & (init_mat_randints(2, 3, 1, 10) <= 10))
```

```python
# grading cell 14

```

## Problem 8: Matrix Initialization
Initialize a matrix of size `m x n` consisting of all elements as `K` and replace all odd values in `arr` with `-1`.

Example:

- Input: 
```
[[5, 8],
 [7, 8],
 [8, 3]]
```
- Output: 
```
[[-1,  8],
 [-1,  8],
 [ 8, -1]]
```

```python
def init_mat_randints(m, n, min_val, max_val):
    # YOUR CODE HERE
    raise NotImplementedError()

def replace_odd_values(matrix):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
matrix = np.array([[5, 8], [7, 8], [8, 3]])
replace_odd_values(matrix)
```

<!-- grade_id: cell-0ef1ef029f07a0d5, points: 5, grade: true, solution: false -->
```python
# grading cell 15
matrix = replace_odd_values(init_mat_randints(4, 3, 1, 10))
assert np.array_equal(matrix, np.where(matrix % 2 != 0, -1, matrix)), "Odd values are not replaced correctly."
```

<!-- grade_id: cell-61b220febe3ab9d1, points: 5, grade: true, solution: false -->
```python
# grading cell 16

```

## Problem 9: Matrix Initialization
Initialize an identity matrix of a specified size.

Example:

- Input:
```
m = 5
```

- Output:
```
[[1, 0, 0, 0, 0]
 [0, 1, 0, 0, 0]
 [0, 0, 1, 0, 0]
 [0, 0, 0, 1, 0]
 [0, 0, 0, 0, 1]]
```

```python
def init_mat_identity_matrix(n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_mat_identity_matrix(5)
```

```python
# grading cell 17
assert init_mat_identity_matrix(3).shape == (3, 3)
assert np.all(np.diag(init_mat_identity_matrix(3)) == 1)
assert np.all(init_mat_identity_matrix(3)[np.where(~np.eye(init_mat_identity_matrix(3).shape[0], dtype=bool))] == 0)
```

```python
# grading cell 18

```

## Problem 10: Matrix Initialization
Initialize a matrix of size `m x n` with elements increasing from left to right and top to bottom. The smallest element starts from `1`.

Example:

- Input:
```
m = 2
n = 4
```

- Output:
```
array([[1, 2, 3, 4],
       [5, 6, 7, 8]])
```

<!-- grade_id: cell-9c8f504ad99b5023, points: 0, grade: false, solution: true -->
```python
def init_mat_increase(m, n):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
# testing with an example
init_mat_increase(2, 4)
```

<!-- grade_id: cell-03844dd4ad0af952, points: 5, grade: true, solution: false -->
```python
# grading cell 19
assert init_mat_increase(2,3).shape == (2, 3)
assert np.sum(abs(init_mat_increase(2,3) - np.array([[1, 2, 3], [4, 5, 6]]))) == 0
```

<!-- grade_id: cell-79c5aaca6a38a24d, points: 5, grade: true, solution: false -->
```python
# grading cell 20

```