# CS414 - Numpy Application
Instructions: 
- Start the kernel: In the menu bar, select Kernel -> Restart kernel.
- Complete all sections with the comment `YOUR CODE HERE`.
- Run all code blocks to check the implementation: In the menu bar, select Cell -> Run All.

```python
import numpy as np
```

## Problem 1: Savings Account

Calculate the remaining balance in the bank account with the input as a list of transaction amounts, corresponding to a list of transaction types indicating deposits (`D`) and withdrawals (`W`).

**Input:**
   - `moneys`: a numpy array containing a list of transaction amounts (deposits and withdrawals)
   - `types`: a numpy array containing a list of transaction types (denoted as `D` for deposit and `W` for withdrawal)

**Output:**
   - A floating-point number representing the remaining balance in the account.

Example:

- Input:
```
moneys = [400.0, 100.0, 200.0]
types = ['D', 'W', 'D']
```

- Output: ``` 500.0 ```

<!-- grade_id: cell-048577ccfe8e62c6, points: 0, grade: false, solution: true -->
```python
def balance(moneys, types):
    # YOUR CODE HERE
    return np.array([moneys[types == 'D'].sum() - moneys[types =='W'].sum()])[0]
    raise NotImplementedError()
```

```python
# testing with an example
balance(np.array([400.0, 100.0, 200.0]), np.array(['D', 'W', 'D']))
```

Output:
```
500.0
```

<!-- grade_id: cell-e03038c165ddaafc, points: 5, grade: true, solution: false -->
```python
# grading cell 1
assert balance(np.array([400.0, 100.0, 200.0]), np.array(['D', 'W', 'D'])) == 500
```

<!-- grade_id: cell-839d4d8b7de04b4c, points: 5, grade: true, solution: false -->
```python
# grading cell 2

```

## Problem 2: GPA
Calculate the average grade of a student given a list of grades and the number of credits for each course.

Input:
   - `marks`: List of length `n`. This represents the grades of the student for each course.
   - `credits`: List of length `n`. This represents the number of credits for each course.

Output:
   - Float value, the average grade of the student's courses.

Example:

- Input:
```
marks = [8.0, 9.0, 10.0]
credits = [2, 2, 1]
```
- Output: ``` 8.8 ```

<!-- grade_id: cell-a0d93b1f727b2f51, points: 0, grade: false, solution: true -->
```python
def gpa_one_student(marks, weights):
    result = 0
    # YOUR CODE HERE
    sum_weighted_grade = np.multiply(marks, weights).sum()
    result = sum_weighted_grade / np.sum(weights)
    return result
    raise NotImplementedError()
```

```python
# testing with an example
marks = [8.0, 9.0, 10.0]
credits = [2, 2, 1]
gpa_one_student(marks, credits)
```

Output:
```
8.8
```

<!-- grade_id: cell-f9ecf7150939ef80, points: 5, grade: true, solution: false -->
```python
# grading cell 3
assert abs(gpa_one_student([10, 9, 8], [4, 3, 2])-9.222222222222221) < 0.000001
```

<!-- grade_id: cell-4706026df476b116, points: 5, grade: true, solution: false -->
```python
# grading cell 4

```

## Problem 3: Student Score Count
Count the number of students in a class who have a score not less than K.

Input:
- `marks`: numpy array of size 1 x n. This represents the math scores of n students in the class.
- `K`: float value, the cutoff score.

Output:
- Integer value, the number of students in the class who have a score equal to or greater than K.

Example:
- Input
```
marks = [4., 9., 5.]
K = 5
```
- Output: ``` 2 ```

<!-- grade_id: cell-75bbb0d6e24b7c40, points: 0, grade: false, solution: true -->
```python
def pass_exam(marks, K):
    # YOUR CODE HERE
    return np.sum(marks >= K)
    raise NotImplementedError()
```

```python
# testing with an example
marks = np.array([4.0, 9.0, 5.0])
K = 5
pass_exam(marks, K)
```

Output:
```
2
```

<!-- grade_id: cell-19c61f2fdfaa710a, points: 5, grade: true, solution: false -->
```python
# grading cell 5
assert pass_exam(np.array([3.0, 9.5, 4.5]), 5.0) == 1
```

<!-- grade_id: cell-6457e32df8fac035, points: 5, grade: true, solution: false -->
```python
# grading cell 6

```

## Problem 4: Highest Class GPA
Calculate the highest grade point average (GPA) of a class given the list of marks and credits for each course.

Input:
   - `marks`: numpy array, size `m x n`. This represents the marks of each student for each course in the class. Here, `m`is the number of students and `n` is the number of courses.
   - `credits`: numpy array, size `1 x n`. This represents the number of credits for each course.

Output:
   - Floating-point number, the highest GPA among the students in the class.

Example:
- Input
```
marks = [[8.0, 9.0, 10.0],
         [7.0, 9.0, 8.0],
         [8.0, 8.0, 8.0],
         [10.0, 9.0, 7.0],
         [9.0, 9.0, 10.0]]
credits = [2, 2, 1]
```
- Output: ``` 9.2 ```

<!-- grade_id: cell-3a099ebd7eee51b7, points: 0, grade: false, solution: true -->
```python
def max_gpa(marks, credits):
    # YOUR CODE HERE
    return np.max(np.divide(np.sum(np.multiply(marks, credits),axis=1), np.sum(credits)))
    raise NotImplementedError()
```

```python
# testing with an example
marks = np.array(
   [[8.0, 9.0, 10.0],
   [7.0, 9.0, 8.0],
   [8.0, 8.0, 8.0],
   [10.0, 9.0, 7.0],
   [9.0, 9.0, 10.0]
])
credits = np.array([2,2,1])
max_gpa(marks, credits)
```

Output:
```
9.2
```

<!-- grade_id: cell-84b3a393314e50ba, points: 5, grade: true, solution: false -->
```python
# grading cell 7
marks = np.array(
   [[4.0, 9.0, 4.5],
   [5.0, 8.0, 4.5],
   [6.0, 7.0, 7.5],
   [7.0, 6.0, 7.5],
   [8.0, 5.0, 4.5]
])
credits = np.array([2,1,2])
assert abs(max_gpa(marks, credits)-7.0) < 0.000001
```

<!-- grade_id: cell-afcc38c40ff1554e, points: 5, grade: true, solution: false -->
```python
# grading cell 8

```

## Problem 5: Cumulative Grade Point Average (CGPA)
Calculate the cumulative grade point average (CGPA) for each student in a class given the list of marks and credits for each course.

**Note:** When calculating the CGPA, we will exclude grades and credits for courses below 5 points (failed courses).

**CGPA Calculation Formula:**  ``` CGPA = (mark1 * credit1 + ... + markn * creditn) / total credits ```

Input:
   - `marks`: numpy array, size `m x n`. This represents the marks of each student for each course in the class. Here, `m` is the number of students and `n` is the number of courses.
   - `credits`: numpy array, size `1 x n`. This represents the number of credits for each course.

Output:
   - numpy array, size `1 x m`, containing the CGPA for each student in the class.

Example:

- Input:
```
marks = [[8.0, 9.0, 10.0],
         [4.0, 9.0, 8.0],
         [8.0, 3.0, 8.0],
         [10.0, 9.0, 5.0],
         [9.0, 9.0, 4.0]
]
credits = np.array([2, 2, 1])
```

- Output:
```
[8.8, 8.66666667, 8., 8.6, 9.]
```

<!-- grade_id: cell-85c2823fb5e72e1e, points: 0, grade: false, solution: true -->
```python
def gpa_of_pass(marks, credits):
    # YOUR CODE HERE
    row_marks, column_marks = marks.shape
    re_mul = np.multiply(np.where(marks >=5, marks, 0),credits)
    cre_dup = np.tile(credits, (row_marks,1))
    return np.divide(np.sum(re_mul, axis = 1), np.sum(np.where(re_mul!=0, cre_dup,0),axis =1))
    raise NotImplementedError()
```

```python
# testing with an example
marks = np.array(
    [[8.0, 9.0, 10.0],
    [4.0, 9.0, 8.0],
    [8.0, 3.0, 8.0],
    [10.0, 9.0, 5.0],
    [9.0, 9.0, 4.0]])
credits = np.array([2,2,1])
gpa_of_pass(marks, credits)
```

Output:
```
array([8.8       , 8.66666667, 8.        , 8.6       , 9.        ])
```

<!-- grade_id: cell-fa2fa11d0d185e58, points: 5, grade: true, solution: false -->
```python
# grading cell 9
marks = np.array(
   [[4.0, 9.0, 4.5],
   [5.0, 8.0, 4.5],
   [6.0, 7.0, 7.5],
   [7.0, 6.0, 7.5],
   [8.0, 5.0, 4.5]])
credits = np.array([2,1,2])
assert np.sum(abs(gpa_of_pass(marks, credits)-np.array([9., 6., 6.8, 7.0, 7.0]))) < 0.000001
```

<!-- grade_id: cell-8119a54a9940d0cf, points: 5, grade: true, solution: false -->
```python
# grading cell 10

```