# CS414 - Object Oriented Programming - Part 1
Instructions: 
- Start the kernel: In the menu bar, select Kernel -> Restart kernel.
- Complete all sections with the comment `YOUR CODE HERE`.
- Run all code blocks to check the implementation: In the menu bar, select Cell -> Run All.

Write a program to implement the following 3 classes: `Geometry`, `Square`, and `Round`. It is known that the `Square` and `Round` classes inherit from the `Geometry` class. 
All 3 classes have the following attributes and methods:

- Attributes:
    - `property` (`float`): can be considered as the side length (for squares) or radius (for circles).

- Methods:
    - Constructor `__init__`.
    - Method to calculate area `calculate_area`.

**Example 1:** the `calculate_area` method in the `Square` class:
- Attribute of the object:
```
property = 2
```
- Input: ``` None ```
- Output: ``` 4 ```

**Example 2:** the `calculate_area` method in the `Round` class:
- Attribute of the object:
```
property = 2
```
- Input: ``` None ```
- Output:
``` 12.56637 ```

**Note:** Use `math.pi` to calculate the area, and round the result to 5 decimal places.

## Problem 1: Implement the `Geometry` class. Implement the `calculate_area` method as an abstract method.


<!-- grade_id: cell-a2a977de6573d46c, points: 0, grade: false, solution: true -->
```python
from abc import abstractmethod

class Geometry:
    def __init__(self, property):
        # YOUR CODE HERE
        raise NotImplementedError()


    @abstractmethod
    def calculate_area(self) -> float:
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-cd8894a81325a585, points: 5, grade: true, solution: false -->
```python
# grading cell 1
geometry_obj = Geometry(2)
assert geometry_obj.property == 2
```

<!-- grade_id: cell-e5f344956ad1440a, points: 5, grade: true, solution: false -->
```python
# grading cell 2

```

## Problem 2:  Implement the `Square` class, inheriting from the `Geometry` class.


<!-- grade_id: cell-735e1d5e4bdbe1e7, points: 0, grade: false, solution: true -->
```python
class Square(Geometry):
    def __init__(self, property):
        # call the `__init__` of the parent class using `super()`
        # YOUR CODE HERE
        raise NotImplementedError()

    # implement `calculate_area`
    def calculate_area(self):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-7d81c0737ff8ff5c, points: 5, grade: true, solution: false -->
```python
# grading cell 3
square_obj = Square(2)
assert square_obj.calculate_area() == 4
```

<!-- grade_id: cell-7483770bd88dec92, points: 5, grade: true, solution: false -->
```python
# grading cell 4

```

## Problems 3: Implement the `Round` class, inheriting from the `Geometry` class. Note the use of `math.pi` and rounding the result to 5 decimal places.

<!-- grade_id: cell-43e8069b1597af6b, points: 0, grade: false, solution: true -->
```python
import math

class Round(Geometry):
    def __init__(self, property):
        # call the `__init__` of the parent class using `super()`
        # YOUR CODE HERE
        raise NotImplementedError()

    # implement `calculate_area`
    def calculate_area(self):
        # YOUR CODE HERE
        raise NotImplementedError()
```

<!-- grade_id: cell-3e614da348102a27, points: 5, grade: true, solution: false -->
```python
# grading cell 5
square_obj = Round(2)
assert square_obj.calculate_area() == 12.56637
```

<!-- grade_id: cell-86742bf27acceec1, points: 5, grade: true, solution: false -->
```python
# grading cell 6

```