# CS414 - Object Oriented Programming - Part 2
Instructions: 
- Start the kernel: In the menu bar, select Kernel -> Restart kernel.
- Complete all sections with the comment `YOUR CODE HERE`.
- Run all code blocks to check the implementation: In the menu bar, select Cell -> Run All.

Write a program to implement the `Fraction` class. The class has the following attributes and methods:

- Attributes:
    - `numerator` (int): Numerator.
    - `denominator` (int): Denominator (must be non-zero, use the `assert` statement to check).

- Methods:
    - Initialization method `__init__`: Raise an error if the denominator is set to `0`. Reduce the fraction if needed.
    - Calculation methods: `addition` (addition), `subtraction` (subtraction):
        - Update the `numerator` and `denominator` attributes of the object. Reduce the fraction if needed.
    - Simplifying a fraction function: `simplify_fraction`.

**Note:** If the numerator is `0`, set the denominator to `1`.

Example 1: the `__init__` method
- Input:
```
numerator = 2
denominator = 4
```
- Output: ``` None ```
- Attributes of the object after execution:
```
numerator = 1
denominator = 2
```

Example 2: the `__init__` method
- Input:
```
numerator = 2
denominator = 0
```
- Output: ``` error ```
- Attributes of the object after execution: ``` error ```

Example 3: the `addition` method (addition of fractions)
- Attributes of the object:
```
numerator = 1
denominator = 3
```
- Input:
```
other = Fraction(1, 3)
```
- Output: ``` None ```
- Attributes of the object after execution:
```
numerator = 2
denominator = 3
```

Example 4: the `subtraction` method:
- Attributes of the object:
```
numerator = 2
denominator = 3
```
- Input:
```
other = Fraction(1, 3)
```
- Output: ``` None ```
- Attributes of the object after execution:
```
numerator = 1
denominator = 3
```

## Exercise: Implement the `Fraction` class.

<!-- grade_id: cell-1b3ae0e4016117ca, points: 0, grade: false, solution: true -->
```python
import math

class Fraction:
    def __init__(self, numerator, denominator):
        # use `assert` to check the denominator
        # YOUR CODE HERE
        assert denominator != 0, ValueError("Denominator cannot be zero")
        assert numerator !=0, self.denominator==1
        
        # implement `numerator`, `denominator` attributes
        # YOUR CODE HERE
        self.numerator = numerator
        self.denominator = denominator
        
        # simplify fraction
        self.simplify_fraction()    

    def simplify_fraction(self):
        # use `math.gcd` to calculate greatest common divisor
        # YOUR CODE HERE
        gcd = math.gcd(self.numerator, self.denominator)
        self.numerator = self.numerator // gcd
        self.denominator = self.denominator // gcd
        print(self.numerator, self.denominator)
        

    def addition(self, other):
        # YOUR CODE HERE
        self.numerator = self.numerator*other.denominator + other.numerator*self.denominator
        self.denominator = self.denominator*other.denominator
        self.simplify_fraction()
        
        
    def subtraction(self, other):
        # YOUR CODE HERE
        self.numerator = self.numerator*other.denominator - other.numerator*self.denominator
        self.denominator = self.denominator*other.denominator
        self.simplify_fraction()
        
```

<!-- grade_id: cell-0a68d569f7353d7e, points: 5, grade: true, solution: false -->
```python
# grading cell 1
fraction_obj = Fraction(2, 4)
assert fraction_obj.numerator == 1
assert fraction_obj.denominator == 2
```

Output:
```
1 2
```

<!-- grade_id: cell-4a46c19c4658a5ea, points: 5, grade: true, solution: false -->
```python
# grading cell 2

```

<!-- grade_id: cell-bfb327985624a91b, points: 5, grade: true, solution: false -->
```python
# grading cell 3
try:
    has_error = False
    Fraction(1, 0)
except:
    has_error = True

assert has_error
```

<!-- grade_id: cell-f867ef309ef64681, points: 5, grade: true, solution: false -->
```python
# grading cell 4
fraction_obj = Fraction(1, 3)
fraction_obj.addition(Fraction(1, 3))
assert fraction_obj.numerator == 2
assert fraction_obj.denominator == 3
```

Output:
```
1 3
1 3
2 3
```

<!-- grade_id: cell-f20b883ad28bf92c, points: 5, grade: true, solution: false -->
```python
# grading cell 5

```

<!-- grade_id: cell-7a92923451041b7f, points: 5, grade: true, solution: false -->
```python
# grading cell 6
fraction_obj = Fraction(2, 3)
fraction_obj.subtraction(Fraction(1, 3))
assert fraction_obj.numerator == 1
assert fraction_obj.denominator == 3
```

Output:
```
2 3
1 3
1 3
```

<!-- grade_id: cell-b326d556e846f458, points: 5, grade: true, solution: false -->
```python
# grading cell 7

```

<!-- grade_id: cell-eda45bc34a10a226, points: 5, grade: true, solution: false -->
```python
# grading cell 8

```